# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    value = True
    xs = str(x)
    ys = str(y)
    
    xl = [i for i in xs]
    yl = [j for j in ys]
    
    xs = sorted(xl)
    ys = sorted(yl)
    
    if len(xs) != len(ys):
        value = False
    else:
        for i in range(0,len(xs)):
            if xs[i]!=ys[i]:
                value = False
    return value
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

Searching for the smallest. So I will look from the bottom and the first one I find is the smallest

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
myint = 0
for b in range(1,200000):
    if same_digits(b, 2*b) and same_digits(b, 3*b) and same_digits(b, 4*b) and same_digits(b, 5*b) and same_digits(b, 6*b)  :
        myint = b
        break
myint



```

Output:
```
142857
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```